cnf(c_0_0, axiom, (q|~p(X1)), file('../../../test.tptp', 1)).
cnf(c_0_1, axiom, (~q), file('../../../test.tptp', 3)).
cnf(c_0_2, axiom, (r(c)|~s(d)), file('../../../test.tptp', 4)).
cnf(c_0_3, axiom, (~r(c)), file('../../../test.tptp', 5)).
cnf(c_0_4, axiom, (p(a)|s(X1)), file('../../../test.tptp', 2)).
cnf(c_0_5, plain, (~p(X1)), inference(sr,[status(thm)],[c_0_0, c_0_1])).
%cnf(c_0_6, plain, (~s(d)), inference(sr,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_7, plain, (s(X1)), inference(sr,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_8, plain, ($false),inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2, c_0_3]),c_0_7])).