cnf(c_0_0, axiom, ~q ).
cnf(c_0_1, axiom, q | p | r).
cnf(c_0_2, axiom, q | ~p).
cnf(c_0_3, axiom, q | ~r).
cnf(c_0_4, plain, ~p , inference(sr,[status(thm)],[c_0_0, c_0_2])).
cnf(c_0_5, plain, ~r , inference(sr,[status(thm)],[c_0_0, c_0_3])).
cnf(c_0_6, plain, p | r , inference(sr,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_7, plain, p , inference(sr,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8, plain, ($false),inference(sr,[status(thm)],[c_0_4, c_0_7])).
